# Part 0: Mining the web

Perhaps the richest source of openly available data today is [the Web](http://www.computerhistory.org/revolution/networking/19/314)! In this lab, you'll explore some of the basic programming tools you need to scrape web data.

> **Warnings.**
> 1. If you are using one of the cloud-based Jupyter installations to run this notebook, such as [Microsoft Azure Notebooks](https://notebooks.azure.com) or [Vocareum](https://vocareum.org), it's likely you will encounter problems due to restrictions on access to remote servers.
> 2. Even if you are using a home or local installation of Jupyter, you may encounter problems if you attempt to access a site too many times or too rapidly. That can happen if your internet service provider (ISP) or the target website detect your accesses as "unusual" and reject them. It's easy to imagine accidentally writing an infinite loop that tries to access a page and being seen from the other side as a malicious program. :)

## The Requests module

Python's [Requests module](http://requests.readthedocs.io/en/latest/user/quickstart/) to download a web page.

For instance, here is a code fragment to download the [Georgia Tech](http://www.gatech.edu) home page and print the first 250 characters. You might also want to [view the source](http://www.computerhope.com/issues/ch000746.htm) of Georgia Tech's home page to get a nicely formatted view, and compare its output to what you see above.

```python
import requests

response = requests.get('http://www.gatech.edu/')
webpage = response.text  # or response.content for raw bytes

print(webpage[0:250]) # Prints the first hundred characters only
```

Output:
```
<!DOCTYPE html PUBLIC "-//W3C//DTD XHTML 1.1//EN" "http://www.w3.org/TR/xhtml11/DTD/xhtml11.dtd">

<html xmlns="http://www.w3.org/1999/xhtml" xml:lang="en">
	<head>
		<title>Test Page for the Apache HTTP Server on Red Hat Enterprise Linux</title>
		<
```

```python
webpage[:10000]
```

Output:
```
'<!DOCTYPE html PUBLIC "-//W3C//DTD XHTML 1.1//EN" "http://www.w3.org/TR/xhtml11/DTD/xhtml11.dtd">\n\n<html xmlns="http://www.w3.org/1999/xhtml" xml:lang="en">\n\t<head>\n\t\t<title>Test Page for the Apache HTTP Server on Red Hat Enterprise Linux</title>\n\t\t<meta http-equiv="Content-Type" content="text/html; charset=UTF-8" />\n\t\t<style type="text/css">\n\t\t\t/*<![CDATA[*/\n\t\t\tbody {\n\t\t\t\tbackground-color: #fff;\n\t\t\t\tcolor: #000;\n\t\t\t\tfont-size: 0.9em;\n\t\t\t\tfont-family: sans-serif,helvetica;\n\t\t\t\tmargin: 0;\n\t\t\t\tpadding: 0;\n\t\t\t}\n\t\t\t:link {\n\t\t\t\tcolor: #c00;\n\t\t\t}\n\t\t\t:visited {\n\t\t\t\tcolor: #c00;\n\t\t\t}\n\t\t\ta:hover {\n\t\t\t\tcolor: #f50;\n\t\t\t}\n\t\t\th1 {\n\t\t\t\ttext-align: center;\n\t\t\t\tmargin: 0;\n\t\t\t\tpadding: 0.6em 2em 0.4em;\n\t\t\t\tbackground-color: #900;\n\t\t\t\tcolor: #fff;\n\t\t\t\tfont-weight: normal;\n\t\t\t\tfont-size: 1.75em;\n\t\t\t\tborder-bottom: 2px solid #000;\n\t\t\t}\n\t\t\th1 strong {\n\t\t\t\tf
```

**Exercise 1.** Given the contents of the GT home page as above, write a function that returns a list of links (URLs) of the "top stories" on the page.

For instance, on Friday, September 9, 2016, here was the front page:

![www.gatech.edu as of Fri Sep 9, 2016](./www.gatech.edu--2016-09-09--annotated-medium.png)

The top stories cycle through in the large image placeholder shown above. We want your function to return the list of URLs behind each of the "Full Story" links, highlighted in red. If no URLs can be found, the function should return an empty list.

<!-- grade_id: get_top_stories, points: 0, grade: false, solution: true -->
```python
import re # Maybe you want to use a regular expression?

def get_gt_top_stories(webpage_text):
    """Given the HTML text for the GT front page, returns a list
    of the URLs of the top stories or an empty list if none are
    found.
    """
#
# YOUR CODE HERE
    site_matcher = re.compile('(www\.\w+\.\w+)')    
    matches = site_matcher.findall(webpage_text)
    return matches

top = get_gt_top_stories(webpage)
top
```

Output:
```
['www.w3.org',
 'www.w3.org',
 'www.example.com',
 'www.redhat.com',
 'www.redhat.com']
```

```python
top_stories = get_gt_top_stories(webpage)
print("Links to GT's top stories:", top_stories)
```

Output:
```
Links to GT's top stories: ['www.w3.org', 'www.w3.org', 'www.example.com', 'www.redhat.com', 'www.redhat.com']
```

## A more complex example

Go to [Yelp!](http://www.yelp.com) and look up `ramen` in `Atlanta, GA`. Take note of the URL:

![Yelp! search for ramen in ATL](./yelp-search-example.png)

This URL encodes what is known as an _HTTP "get"_ method (or request). It basically means a URL with two parts: a _command_ followed by one or more _arguments_. In this case, the command is everything up to and including the word `search`; the arguments are the rest, where individual arguments are separated by the `&` or `#`.

> "HTTP" stands for "HyperText Transport Protocol," which is a standardized set of communication protocols that allow _web clients_, like your web browser or your Python program, to communicate with _web servers_.

In this next example, let's see how to build a "get request" with the `requests` module. It's pretty easy!

```python
url_command = 'http://yelp.com/search'
url_args = {'find_desc': "ramen",
            'find_loc': "atlanta, ga"}
response = requests.get (url_command, params=url_args)

# print ("==> Downloading from: '%s'" % response.url) # confirm URL
# print ("\n==> Excerpt from this URL:\n\n%s\n" % response.text)
yelp_search = response.text
yelp_search[:100000]
```

Output:
```
'<!DOCTYPE HTML>\n\n<!--[if lt IE 7 ]> <html xmlns:fb="http://www.facebook.com/2008/fbml" class="ie6 ie ltie9 ltie8 no-js" lang="en"> <![endif]-->\n<!--[if IE 7 ]>    <html xmlns:fb="http://www.facebook.com/2008/fbml" class="ie7 ie ltie9 ltie8 no-js" lang="en"> <![endif]-->\n<!--[if IE 8 ]>    <html xmlns:fb="http://www.facebook.com/2008/fbml" class="ie8 ie ltie9 no-js" lang="en"> <![endif]-->\n<!--[if IE 9 ]>    <html xmlns:fb="http://www.facebook.com/2008/fbml" class="ie9 ie no-js" lang="en"> <![endif]-->\n<!--[if (gt IE 9)|!(IE)]><!--> <html xmlns:fb="http://www.facebook.com/2008/fbml" class="no-js" lang="en"> <!--<![endif]-->\n    <head>\n\n\n    <script>            window.yPageStart = new Date().getTime();\n</script>\n\n    <script>            var initialVisibilityState = document.webkitVisibilityState;\n\n                yPerfTimings = [];\n\n                ySitRepParams = {"b3Sampled": null, "edgeStartTime": 1508719565986634, "site": "main", "edgestageSubdomain": "_disabled", "
```

**Exercise 2.** Given a search topic, location, and a rank $k$, return the name of the $k$-th item of a Yelp! search. If there is no $k$-th item, return `None`.

> The demo query above only gives you a website with the top 10 items, meaning you could only use it for $k \leq 10$. Figure out how to modify it to solve the problem when $k > 10$.

<!-- grade_id: yelp_find_item, points: 0, grade: false, solution: true -->
```python
def find_yelp_item (topic, location, k):
    """Returns the k-th suggested item from Yelp! in Atlanta for the given topic."""
#
# YOUR CODE HERE
#
    url_command = 'http://yelp.com/search'
    url_args = {'find_desc': topic, 'find_loc': location, 'start': k-1}
    response = requests.get(url_command, params=url_args)
    yelp_search = response.text
    
    from bs4 import BeautifulSoup
    import re

    yelp_soup = BeautifulSoup(yelp_search, 'html.parser')
    html_business = yelp_soup.findAll('span', attrs={'class':"indexed-biz-name"})

    busi_list = []
    for result in html_business:
        business = result.find('a', attrs={'class': "biz-name js-analytics-click"}).span.get_text()
        busi_list.append(business)
    return(busi_list[0])

find_yelp_item('ramen', 'atlanta, ga', 17)
```

Output:
```
'Pijiu Belly'
```

<!-- grade_id: yelp_atl__test1, points: 0, grade: true, solution: false -->
```python
assert find_yelp_item('fried chicken', 'Atlanta, GA', -1) is None # Tests an invalid value for 'k'
```

Output:
```
<ipython-input-58-db097f8795d2> in <module>()
----> 1 assert find_yelp_item('fried chicken', 'Atlanta, GA', -1) is None # Tests an invalid value for 'k'

<ipython-input-53-f5bc1743b3d0> in find_yelp_item(topic, location, k)
     19         business = result.find('a', attrs={'class': "biz-name js-analytics-click"}).span.get_text()
     20         busi_list.append(business)
---> 21     return(busi_list[0])
     22 
     23 find_yelp_item('ramen', 'atlanta, ga', 17)

IndexError: list index out of range
```

> Search queries on Yelp! don't always return the same answers, since the site is always changing! Also, your results might not match a query you do via your web browser (_why not?_). As such, you should manually check your answers.

<!-- grade_id: yelp_atl__test2, points: 0, grade: true, solution: false -->
```python
item = find_yelp_item ('fried chicken', 'Atlanta, GA', 1)
print (item)

# The most likely answer on September 19, 2017:
#assert item in ['Guss World Famous <span class="highlighted">Fried</span> <span class="highlighted">Chicken</span>',
#                'Guss World Famous Fried Chicken']                
```

Output:
```
Guss World Famous Fried Chicken
```

<!-- grade_id: yelp_atl__test3, points: 0, grade: true, solution: false -->
```python
item = find_yelp_item ('fried chicken', 'Atlanta, GA', 5)
print (item)

# The most likely answer on September 19, 2017:
#assert item == 'Richards Southern Fried'
```

Output:
```
Richards Southern Fried
```

<!-- grade_id: yelp_atl__test4, points: 0, grade: true, solution: false -->
```python
item = find_yelp_item('fried chicken', 'Atlanta, GA', 17)
print(item)

# Most likely correct answer as of September 19, 2017:
#assert item == 'Sway'
```

Output:
```
Harolds Chicken & Ice Bar
```

One issue with the above exercises is that they treat HTML as a flat string, whereas the document is at least semi-structured. Moreover, web pages are such a common source of data today that you would expect better tools for processing them. Indeed, such tools exist! The next part of this assignment, Part 1, walks you through one such tool. So, head there when you are ready!